# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


```python
import pandas as pd
```

```python
df_education = pd.read_csv(r'assets\NISPUF17.csv')
```

```python
columns_field = df_education.columns
```

```python
'EDUC1' in columns_field
```

Output:
```
True
```

```python
df_education.shape
```

Output:
```
(28465, 454)
```

SEQNUMC – unique child ID variable <br>
SEQNUMHH – unique household ID variable <br>
EDUC1 – education of the mother <br>


EDUC1 – education of the mother <br>
0 : <12 years
1 : 12 years
2 : >12 years, not a college graduate
3 : College graduate

```python

```

```python
df_education.EDUC1.value_counts()
```

Output:
```
4    13656
3     6999
2     4906
1     2904
Name: EDUC1, dtype: int64
```

```python
less_hs = df_education.EDUC1.value_counts()[1]/df_education.shape[0]
less_hs

```

Output:
```
0.10202002459160373
```

```python
list_keys = ["less than high school","high school","more than high school but not college","college"]
dict_final = {}
for i in range(0,len(list_keys)):
    dict_final[list_keys[i]] = df_education.EDUC1.value_counts()[i+1]/df_education.shape[0]

```

```python
dict_final
```

Output:
```
{'less than high school': 0.10202002459160373,
 'high school': 0.172352011241876,
 'more than high school but not college': 0.24588090637625154,
 'college': 0.47974705779026877}
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    list_keys = ["less than high school","high school","more than high school but not college","college"]
    dict_final = {}
    for i in range(0,len(list_keys)):
        dict_final[list_keys[i]] = df_education.EDUC1.value_counts()[i+1]/df_education.shape[0]
    return dict_final
    # raise NotImplementedError()
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

CBF_01 – child ever fed breast milk
Yes
No
Don’t know
Missing

P_NUMFLU – total number of seasonal influenza doses

```python
'P_NUMFLU' in df_education
```

Output:
```
True
```

```python
'CBF_01'in df_education
```

Output:
```
True
```

```python
df_education.P_NUMFLU.value_counts()
```

Output:
```
3.0    4845
0.0    4139
2.0    3350
1.0    1679
4.0    1282
5.0      35
6.0       3
Name: P_NUMFLU, dtype: int64
```

being fed breastmilk as a child : df_education.CBF_01 = 1 
getting a seasonal influenza vaccine df_education.P_NUMFLU > 0

```python
avg_vacc_bm_yes = df_education[(df_education.CBF_01 == 1) & (df_education.P_NUMFLU >= 1)].P_NUMFLU.mean()
avg_vacc_bm_not = df_education[(df_education.CBF_01 == 2) & (df_education.P_NUMFLU >= 1)].P_NUMFLU.mean()
```

```python
(avg_vacc_bm_yes,avg_vacc_bm_not)
```

Output:
```
(2.5431043256997454, 2.377330350484713)
```

```python
df_education.CBF_01.value_counts()
```

Output:
```
1     24261
2      4085
77      110
99        9
Name: CBF_01, dtype: int64
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    avg_vacc_bm_yes = df_education[(df_education.CBF_01 == 1) & (df_education.P_NUMFLU >= 1)].P_NUMFLU.mean()
    avg_vacc_bm_not = df_education[(df_education.CBF_01 == 2) & (df_education.P_NUMFLU >= 1)].P_NUMFLU.mean()
    return (avg_vacc_bm_yes,avg_vacc_bm_not)
    raise NotImplementedError()
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

HAD_CPOX – did child ever have chicken pox 
1 Yes
2 No
Don’t know
Refused
Missing
P_NUMVRC – total number of varicella doses
SEX – sex of child
1 Male
2 Female

```python
df_education.HAD_CPOX.value_counts()
```

Output:
```
2     27955
1       402
77      105
99        3
Name: HAD_CPOX, dtype: int64
```

```python
df_education.P_NUMVRC.value_counts()
```

Output:
```
1.0    13781
0.0     1338
2.0      213
3.0        1
Name: P_NUMVRC, dtype: int64
```

```python
df_education.SEX.value_counts()[1]
```

Output:
```
14717
```

```python
total_male = df_education.SEX.value_counts()[1]
total_female = df_education.SEX.value_counts()[2]
```

```python
# number of children who contracted chickenpox  (HAD_CPOX = 1) but were vaccinated (P_NUMVRC>0)
male_chickenpox = df_education[(df_education.HAD_CPOX == 1)&(df_education.P_NUMVRC>0) & (df_education.SEX ==1)].shape[0] / total_male
female_chickenpox =  df_education[(df_education.HAD_CPOX == 1)&(df_education.P_NUMVRC>0) & (df_education.SEX ==2)].shape[0] / total_female

with_cp = {
    'male' : male_chickenpox,
    'female' : female_chickenpox
}
```

```python
# number of children who not contracted chickenpox  (HAD_CPOX = 2) but were vaccinated (P_NUMVRC>0)
male_not_cp = df_education[(df_education.HAD_CPOX == 2)&(df_education.P_NUMVRC>0) & (df_education.SEX ==1)].shape[0] / total_male
female_not_cp =  df_education[(df_education.HAD_CPOX == 2)&(df_education.P_NUMVRC>0) & (df_education.SEX ==2)].shape[0] / total_female

not_cp = {
    'male' : male_not_cp,
    'female' : female_not_cp
}
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    male_not_cp = df_education[(df_education.HAD_CPOX == 2)&(df_education.P_NUMVRC>0) & (df_education.SEX ==1)].shape[0] / total_male
    female_not_cp =  df_education[(df_education.HAD_CPOX == 2)&(df_education.P_NUMVRC>0) & (df_education.SEX ==2)].shape[0] / total_female

    not_cp = {
        'male' : male_not_cp,
        'female' : female_not_cp
    }
    return not_cp
    raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. If the `had_chickenpox_column` is either `1` (for yes) or `2` for no, and that the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine, then a positive correlation (e.g. `corr > 0`) would mean that an increase in `had_chickenpox_column` (which means more no's) would mean an increase in the `num_chickenpox_vaccine_column` (which means more doses of vaccine). If `corr < 0` then there is a negative correlation, indicating that having had chickenpox is related to an increase in the number of vaccine doses. Also, `pval` refers to the probability the relationship observed is significant. In this case `pval` should be very very small (will end in `e-18` indicating a very small number), which means the result unlikely to be by chance.

[1] This isn't really the full picture, since we are not looking at when the dose was given. It's possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

```python
had_cp_vacc = df_education[['HAD_CPOX','P_NUMVRC']]

```

```python
had_cp_vacc.dropna(inplace=True)
```

Output:
```
C:\Users\vmena\AppData\Local\Temp\ipykernel_23092\2106075192.py:1: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  had_cp_vacc.dropna(inplace=True)
```

```python

```

Output:
```
2     15113
1       173
77       47
Name: HAD_CPOX, dtype: int64
```

```python
had_cp_vacc.P_NUMVRC.value_counts()
```

Output:
```
1.0    13781
0.0     1338
2.0      213
3.0        1
Name: P_NUMVRC, dtype: int64
```

```python
had_cp_vacc2 = had_cp_vacc[(had_cp_vacc.HAD_CPOX<=2) & (had_cp_vacc.P_NUMVRC>0.0)]

```

```python
had_cp_vacc2.HAD_CPOX.value_counts()
```

Output:
```
2    13830
1      121
Name: HAD_CPOX, dtype: int64
```

```python
had_cp_vacc2.P_NUMVRC.value_counts()
```

Output:
```
1.0    13737
2.0      213
3.0        1
Name: P_NUMVRC, dtype: int64
```

#1

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    # df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
    #                "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})
    # df=pd.DataFrame({"had_chickenpox_column":had_cp_vacc.HAD_CPOX,
    #             "num_chickenpox_vaccine_column":had_cp_vacc.P_NUMVRC})


    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(had_cp_vacc2.HAD_CPOX,had_cp_vacc2.P_NUMVRC)
    
    # just return the correlation
    return corr

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python
corr_chickenpox()
```

Output:
```
-0.013337081886400624
```

```python
corr_df = had_cp_vacc2.corr()
corr_df
```

Output:
```
HAD_CPOX  P_NUMVRC
HAD_CPOX  1.000000 -0.013337
P_NUMVRC -0.013337  1.000000
```

```python

```